# Math 187B - Homework 1
## Due April 9, 2021 at 9pm


## Rules:


##### <font color="red"> 1. Write your solutions in this file. Do not rename it!<br> 2. Write your answers in the cells indicated for that purpose. Delete the two (autogenerated) lines</font>
    # YOUR CODE HERE
    raise NotImplementedError()
 ##### <font color="red">   and write your solution there. <br> 3. Do not forget to evaluate each cell (Shift + Enter). <br>  4. In order to get any credit, your code has to actually run. If it gives a warning/error, you will receive no credit whatsoever for that (part of the) problem.<br> 5. It is imperative that you name your functions and variables as asked in the problems. No credit will be given otherwise.  <br> 6. Make sure your answers pass the tests given for each problem by evaluating the test cells (Shift + Enter).</font>

### Problem 1: Flow control

Write a function named **fizz_buzz** that takes as input an integer **N** and for each integer **m** from 1 to **N**, and does two things:
 - prints 'Fizz' if **m** is divisible by 2 and not 5, prints 'Buzz' if **m** is divisible by 5 and not 2, prints 'FizzBuzz' if **m** is divisible by 2 and 5, prints 'Moot' if none of the above;
 - returns as output a list of these print statements

eg.
Calling fizz_buzz(10)

should produce
```
Moot
Fizz
Moot
Fizz
Buzz
Fizz
Moot
Fizz
Moot
FizzBuzz
['Moot', 'Fizz', 'Moot', 'Fizz', 'Buzz', 'Fizz', 'Moot', 'Fizz', 'Moot', 'FizzBuzz']
```

<!-- grade_id: 1b0cc4, points: 0, grade: false, solution: true -->
```python
#1 Solution
def fizz_buzz(N):
    output = []
    if type(N) != int:
        output.append('N is not an int')
        return output
    for m in range(1, N+1):
        if m%2 == 0 and m%5 != 0:
            print('Fizz')
            output.append('Fizz')
        elif m%2 != 0 and m%5 == 0:
            print('Buzz')
            output.append('Buzz')
        elif m%2 == 0 and m%5 == 0:
            print('FizzBuzz')
            output.append('FizzBuzz')
        else:
            print('Moot')
            output.append('Moot')
    return output
```

```python
fizz_buzz(3.5)
```

Output:
```
['N is not an int']
```

<!-- grade_id: 849bf1, points: 4, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(fizz_buzz(2) == ['Moot','Fizz'])
assert(fizz_buzz(10) == ['Moot', 'Fizz', 'Moot', 'Fizz', 'Buzz', 'Fizz', 'Moot', 'Fizz', 'Moot', 'FizzBuzz'])
assert(type(fizz_buzz(3) == list))
assert(type(fizz_buzz(3)[0]) == str)
assert(fizz_buzz(3) == ['Moot','Fizz','Moot'])


```

Output:
```
Moot
Fizz
Moot
Fizz
Moot
Fizz
Buzz
Fizz
Moot
Fizz
Moot
FizzBuzz
Moot
Fizz
Moot
Moot
Fizz
Moot
Moot
Fizz
Moot
```

# Problem 2: Strings to numbers

This problem will encode a way to go between strings of characters and integers. Since now we use modern encryption techniques, we want to keep all the punctuation. However this means that we want to treat apostrophes and quotes as punctuation and not delimiters of strings. One way around this is to input all the strings using triple double quotes, e.g. """A CAT!""".

a. Write a function **str2num(s)** that takes as input a string of characters, capitalizes it and returns the SAGE integer obtained by concatenating the corresponding  ASCII codes, e.g. str2num("""A cat!""") returns $653267658433$ (a SAGE integer, not a string and not a Python integer!). (Note that capitalization ensures all ASCII codes are two digit numbers. You may assume all punctuations in the input also have ASCII codes being two digit numbers so don't strip any of them.)

b. Write a function **num2str(m)** that is the inverse of the function in (a), takes as input an integer and returns the corresponding message (capitalized), e.g. num2str(653267658433) returns 'A CAT!' or "A CAT!"

c. Write a function **remove_punctuation(s)** that takes as input a string **s** and returns the same string stripped of spaces and punctuation, and converts all letters to upper case. It should return the error 'Invalid input' if s is not a string, or if s contains anything other than letters, spaces and punctuation; for our purposes punctuation will be the things listed at the top of: https://simple.wikipedia.org/wiki/Punctuation
The goal here it to obtain an output that can be Caesar shifted.

d. Write a function **str26(s)** that takes as input a string **s**, computes **t = remove_punctuation(s)** , and returns a list of integers between 0 and 25 corresponding to each character of **t** in order, where $A\to 0, B\to 1, \dots , Z\to 25.$ If the string **s** contains any digits or special characters, it should return the message 'Invalid input'.

<!-- grade_id: 6e555f, points: 0, grade: false, solution: true -->
```python
# 2a Solution:
def str2num(s):
    s = s.upper()
    output = ''
    for char in s:
        num = ord(char)
        output = output + str(num)
    return int(output)
```

```python
str2num('Jordi (')
```

Output:
```
74798268733240
```

<!-- grade_id: c33b90, points: 2, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(str2num("A cat!") == 653267658433)
assert(str2num("Howdy, howdy!") == 72798768894432727987688933)
assert(str2num("""Just because it has always been done that way doesn't mean it's not incredibly stupid.""") == 7485838432666967658583693273843272658332657687658983326669697832687978693284726584328765893268796983783984327769657832738439833278798432737867826968736676893283848580736846)
assert(str2num("""Never say "never". """) == 78698669823283658932347869866982344632)

```

<!-- grade_id: 62fc69, points: 0, grade: false, solution: true -->
```python
# 2b Solution:
def num2str(m):
    m = str(m)
    output = ''
    for i in range(0, len(m), 2):
        char = chr(int(m[i:i+2]))
        output = output + char
    return output
```

```python
num2str(6565)
```

Output:
```
'AA'
```

<!-- grade_id: d6241a, points: 2, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(num2str(653267658433) == 'A CAT!')
assert(num2str(72798768894432727987688933) == 'HOWDY, HOWDY!')
assert(num2str(7485838432666967658583693273843272658332657687658983326669697832687978693284726584328765893268796983783984327769657832738439833278798432737867826968736676893283848580736846) == """JUST BECAUSE IT HAS ALWAYS BEEN DONE THAT WAY DOESN'T MEAN IT'S NOT INCREDIBLY STUPID.""")


```

<!-- grade_id: cf8989, points: 0, grade: false, solution: true -->
```python
# 2c Solution:
import string
def remove_punctuation(s):
    output = ''
    if type(s) != str:
        output = 'Invalid input'
        return output
    s = s.upper()
    alphabet = string.ascii_uppercase
    l = list(alphabet)
    l.extend(['.', ',', '!', '?', ' ', '"', ':', ';', '...', '-', '_', '–', '(', ')', '[', ']', '\''])
    for char in s:
        if char not in l:
            return 'Invalid input'
        else:
            if char.isalnum():
                output = output + char
    return output
```

```python
remove_punctuation("""Just bec>ause it has always been done that way doesn't mean it's not incredibly stupid.""")
```

Output:
```
'Invalid input'
```

<!-- grade_id: e2591b, points: 2, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(remove_punctuation("A cat!") == 'ACAT')
assert(remove_punctuation("Howdy, <hello> howdy!") == 'Invalid input')
assert(remove_punctuation("""Just because it has always been done that way doesn't mean it's not incredibly stupid.""") == 'JUSTBECAUSEITHASALWAYSBEENDONETHATWAYDOESNTMEANITSNOTINCREDIBLYSTUPID')

```

<!-- grade_id: 075b76, points: 0, grade: false, solution: true -->
```python
# 2d Solution:
def str26(s):
    output = []
    t = remove_punctuation(s)
    if t == 'Invalid input':
        return 'Invalid input'
    for char in t:
        num = ord(char)
        num = num - 65
        output.append(num)
    return output
```

```python
str26("Howdy, <oh jeebers> howdy!")
```

Output:
```
'Invalid input'
```

<!-- grade_id: 73c860, points: 2, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:


assert(str26("A cat!") == [0, 2, 0, 19])
assert(str26("Howdy, <oh jeebers> howdy!") == 'Invalid input')
assert(str26("""Just because it has always been done that way doesn't mean it's not incredibly stupid.""") == [9, 20, 18, 19, 1, 4, 2, 0, 20, 18, 4, 8, 19, 7, 0, 18, 0, 11, 22, 0, 24, 18, 1, 4, 4, 13, 3, 14, 13, 4, 19, 7, 0, 19, 22, 0, 24, 3, 14, 4, 18, 13, 19, 12, 4, 0, 13, 8, 19, 18, 13, 14, 19, 8, 13, 2, 17, 4, 3, 8, 1, 11, 24, 18, 19, 20, 15, 8, 3])

```

### Problem 3: Caesar shifts

Again, some of this is very similar to a problem from HRead about Caesar shifts at https://en.wikipedia.org/wiki/Caesar_cipher


a. Write a function **ngram(s,N)** that takes as input a string **s**, and a positive integer **N**; returns the string **s** separated into N-grams, all in capital letters and stripped of punctuation. For instace ngram('A cat!', 3) should return 'ACA T'. If **N** is not a positive integer or **s** is not a string, or if s contains numbers or special characters, it should return the message 'Invalid input'.



b. Write a function **encrypt_caesar_gram (s,k,N)** that takes as input a string **s**, a key **k** that consists of a single letter, and a positive integer **N**; returns the **k** Caesar shift of the string **s** stripped of spaces and punctuation into **N**-grams and written as upper case letters. If **s** is not a string, if **s** contains special characters, if **k** is not a single letter or if **N** is not a positive integer, it should return the message 'Invalid input'. That is, if you could not do this by hand, then your function should return the error message. 


c. Write a function **decrypt_caesar(s,k)** that takes as input a string **s**, and a key **k** that consists of a single letter; applies **remove_punctuation(s)**, and then decrypts the Caesar shift given by **k** on the stripped string **remove_punctuation(s)**. The output should be a string with no spaces and containing only upper case letters. If the string **s** is not of the expected form (e.g. it contains digits or special characters), it should return the error message 'Invalid input'. If the key **k** is anything other than a single letter, it should return the error message 'Invalid input'. (You might find the functions from Problem 1 useful.)

<!-- grade_id: 949a28, points: 0, grade: false, solution: true -->
```python
# 3a Solution:
def ngram(s,N):
    if N <= 0:
        return 'Invalid input'
    t = remove_punctuation(s)
    if t == 'Invalid input':
        return 'Invalid input'
    t = ' '.join(t[i: i+N] for i in range(0, len(t), N))
    return t
```

```python
ngram("A cat!",3)
```

Output:
```
'ACA T'
```

<!-- grade_id: 25205a, points: 2, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(ngram("A cat!",3) == 'ACA T')
assert(ngram("Howdy <hello> howdy!",4) == 'Invalid input')
assert(ngram("What a novel idea!",0) == 'Invalid input')

```

<!-- grade_id: a0c62b, points: 0, grade: false, solution: true -->
```python
# 3b Solution:
def encrypt_caesar_gram(s,k,N):
    s = ngram(s,N)
    if s == 'Invalid input':
        return 'Invalid input'
    if len(k) > 1 or len(k) == 0:
        return 'Invalid input'
    original = str26(s)
    t = str26(k)
    if t == 'Invalid input':
        return 'Invalid input'
    leng = len(original)
    encrypt = [None]* leng
    for i in range(0, leng):
        num = (original[i] + t[0])%26
        encrypt[i] = num
    output = ''
    for val in encrypt:
        asc = chr(val+65)
        output = output + asc
    output = ngram(output, N)
    return output
```

```python
encrypt_caesar_gram("A cat!", "C", 3)
```

Output:
```
'CEC V'
```

```python
encrypt_caesar_gram("Here is an example","E", 5)
```

Output:
```
'LIVIM WERIB EQTPI'
```

<!-- grade_id: d9f479, points: 3, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(encrypt_caesar_gram("A cat!","B",3) == 'BDB U')
assert(encrypt_caesar_gram("Math is fun & cool!",2,6) == 'Invalid input')
assert(encrypt_caesar_gram("This is spooky","K",-1) == 'Invalid input')


```

<!-- grade_id: 57c64e, points: 0, grade: false, solution: true -->
```python
# 3c Solution:
def decrypt_caesar(s,k):
    s = remove_punctuation(s)
    if s == 'Invalid input':
        return 'Invalid input'
    s = s.upper()
    if len(k) > 1 or len(k) == 0:
        return 'Invalid input'
    if s == 'Invalid input':
        return 'Invalid input'
    t = str26(k)
    if t == 'Invalid input':
        return 'Invalid input'
    original = str26(s)
    leng = len(original)
    decrypt = [None]* leng
    for i in range(0, leng):
        num = (original[i] - t[0])%26
        decrypt[i] = num
    output = ''
    for val in decrypt:
        asc = chr(val+65)
        output = output + asc
    return output
```

```python
decrypt_caesar(12, 'D')
```

Output:
```
'Invalid input'
```

<!-- grade_id: 20c7b7, points: 3, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(decrypt_caesar('XLMW MWEJ YRXM QI',"E") == 'THISISAFUNTIME')
assert(decrypt_caesar("Hello","OK") == 'Invalid input')
assert(decrypt_caesar(12,"A") == 'Invalid input')


```

### Problem 4: Vigenere

Read about the Vigenere cipher at https://en.wikipedia.org/wiki/Vigen%C3%A8re_cipher
We will only use the concepts under **Description**.

a. Write a function **encrypt_vigenere (s, k, N)** that takes as input two strings **s** and **k** and a positive integer **N**; removes the spaces and punctuation from both of them; and then encrypts the stripped string **s** by a Vigenere cipher with the key given by the stripped string **k** written as **N**-grams of capital letters. If either **s** or **k** is not a string or contain digits or special characters, or if **k** is the empty string, it should return the message 'Invalid input'. Same if **N** is not a positive integer.

b. Write a function **decrypt_vigenere (s, k)** that takes as input two strings **s** and **k**; removes the spaces and punctuation from both of them; and then performs the Vigenere decryption of the stripped string **s** assuming the encrytion key used was the stripped string **k**. The output should be a string consisting only of capital letters. If either **s** or **k** contain digits or special characters, it should return the message 'Invalid input'. Same if **N** is not a positive integer.

<!-- grade_id: ca1ecf, points: 0, grade: false, solution: true -->
```python
# 4a Solution:
def encrypt_vigenere(s,k,N):
    s = remove_punctuation(s)
    if s == 'Invalid input':
        return 'Invalid input'
    k = remove_punctuation(k)
    if k == 'Invalid input':
        return 'Invalid input'
    if N <= 0:
        return 'Invalid input'
    t = str26(s)
    if t == 'Invalid input': return t
    if len(k) == 0: return 'Invalid input'
    if str26(k) == 'Invalid input': return 'Invalid input'
    key = (str26(k)*len(s))[:len(s)]
    output = ''
    for i in range(len(t)):
        output += chr(((t[i]+key[i])%26)+65)
    output = ngram(output, N)
    return output
```

<!-- grade_id: 010fb9, points: 4, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(encrypt_vigenere("Math is fun and cool!","ORANGE",4) == 'ARTU OWTL NNTH QFOY')
assert(encrypt_vigenere("What a wonderful word","key",-1) == 'Invalid input')
assert(encrypt_vigenere("What a great <sentence>","lemon",2) == 'Invalid input')

```

```python
# 4b Solution:
### BEGIN SOLUTION

def decrypt_vigenere(s,k):
    t = str26(s)
    if t == 'Invalid input': return t
    if len(k) == 0: return 'Invalid input'
    if str26(k) == 'Invalid input': return 'Invalid input'
    key = (str26(k)*len(s))[:len(s)]
    output = ''
    for i in range(len(t)):
        output += chr(((t[i]-key[i])%26)+65)
    return output

### END SOLUTION
```

<!-- grade_id: 5b579f, points: 4, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(decrypt_vigenere('ARTU OWTL NNTH QFOY',"ORANGE") == 'MATHISFUNANDCOOL')
assert(decrypt_vigenere('G L Y D E U Y R B O V D E P U Y V J N',"key") == 'WHATAWONDERFULWORLD')
assert(decrypt_vigenere("What a great <sentence>","lemon") == 'Invalid input')


```

## Problem 5 (Bonus): Fibonacci numbers

In this problem you have to use the recurrence formula for Fibonacci numbers $f_n = f_{n-1} + f_{n-2}, \quad f_1 = f_2 =1.$

a. Write an _efficient_ function **fib(N)** that takes input a positive integer **N** and returns the Nth Fibonacci number.

b. Now use this to write another function **allfib(N)** that takes input a positive integer **N** and returns a list of the first N Fibonacci numbers.

Warning: your code should be efficient enough so that all tests take less than 30 seconds.
Hint: it might be easier to solve the second part first.

<!-- grade_id: 8b2e77, points: 0, grade: false, solution: true -->
```python
#5a Solution

# YOUR CODE HERE
def fib(N):
    list = [None] * (N+2)
    list[0] = 1
    list[1] = 1
    for i in range(2, N):
        list[i] = list[i-1] + list[i-2]
    return list[N-1]
```

```python
fib(1000)
```

Output:
```
43466557686937456435688527675040625802564660517371780402481729089536555417949051890403879840079255169295922593080322634775209689623239873322471161642996440906533187938298969649928516003704476137795166849228875
```

<!-- grade_id: 803455, points: 5, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(fib(1) == 1)

assert(fib(2) == 1)
assert(fib(3) == 2)
assert(fib(4) == 3)
assert(fib(5) == 5)
assert(fib(6) == 8)
assert(fib(100) == 354224848179261915075)

```

<!-- grade_id: 05c6a0, points: 0, grade: false, solution: true -->
```python
#5b Solution
def allfib(N):
    if N == 1:
        return [1]
    list = [None] * (N+2)
    list[0] = 1
    list[1] = 1
    for i in range(2, N):
        list[i] = list[i-1] + list[i-2]
    result = []
    for num in list:
        if num != None:
            result.append(num)
    return result
```

```python
allfib(10)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13, 21, 34, 55]
```

<!-- grade_id: 5e7037, points: 5, grade: true, solution: false -->
```python
# Check that your function returns the correct output for several inputs:

assert(allfib(1) == [1])
assert(allfib(2) == [1,1])
assert(allfib(3) == [1,1,2])
assert(allfib(8) == [1, 1, 2, 3, 5, 8, 13, 21])
assert(allfib(100)[-1] == 354224848179261915075)
assert(allfib(100)[-1] == fib(100) )
assert(len(allfib(10)) == 10)
assert(type(allfib(10)) == list)
assert(type(allfib(10)[0])== int)

```

```python

```

```python

```

```python

```

```python

```

```python

```